# Homework 1


## Problem 1 (If-statment)

If statement:

Given an integer n, perform the following conditional actions:

    If n is odd, set output to "Hello World"
    If n is even and in the inclusive range of 2 to 5, set output to "Goodbye World"
    If n is even and in the inclusive range of 6 to 20, set output to "Hello World"
    If n is even and greater than 20, set output to "Goodbye World"

Example:

    -- Input: n = 2
    -- Output: output = "Goodbye World"

Write your code here

<!-- grade_id: cell-b78086c290d4a60c, points: 0, grade: false, solution: true -->
```python
def problem1(n):
    output = 0
    # YOUR CODE HERE 
    if n % 2 != 0 or 6 <= n <= 20  :
        output = "Hello World"
        print(f"out put = {output}")
    elif 2 <= n <= 5 or 20 < n:
        output = "Goodbye World"
        print(f"out put = {output}")
    else:
        raise NotImplementedError()
    return(output)

```

<!-- grade_id: cell-2853176972493b6e, points: 25, grade: true, solution: false -->
```python
assert problem1(3) == "Hello World"
assert problem1(2) == "Goodbye World"
assert problem1(18) == "Hello World"
```

Output:
```
out put = Hello World
out put = Goodbye World
out put = Hello World
```

<!-- grade_id: cell-b16dd3056b5a9f95, points: 25, grade: true, solution: false -->
```python
assert problem1(29) == "Hello World"
assert problem1(24) == "Goodbye World"
```

Output:
```
out put = Hello World
out put = Goodbye World
```

Write ChatGPT code here

```python
def problem1(n):
    if n % 2 != 0:  # 如果 n 是奇数
        return "Hello World"
    elif 2 <= n <= 5:  # 如果 n 是偶数并且在 2 到 5 的范围内
        return "Goodbye World"
    elif 6 <= n <= 20:  # 如果 n 是偶数并且在 6 到 20 的范围内
        return "Hello World"
    elif n > 20:  # 如果 n 是偶数并且大于 20
        return "Goodbye World"

```

```python
assert problem1(3) == "Hello World"
assert problem1(2) == "Goodbye World"
assert problem1(18) == "Hello World"
```

```python
assert problem1(29) == "Hello World"
assert problem1(24) == "Goodbye World"
```

## Problem 2 (For-loop, basic list manipulation)

You and $K-1$ friends want to buy $N$ different pieces of candy.  Candy piece $i$ costs $p_i$. The candy shop owner is greedy and wants to maximize his number of new customers, so he increases the sale price of the candy for repeated customers; more precisely, if a customer has already purchased $x$ pieces of candy, then the cost of candy $i$ is $C_i*(x+1)$.

Find and print the minimum cost for your group to purchase all $N$ pieces of candy.

The input:

    N: the number of candies to buy
    K: the number of people
    C: the cost of candy

The output:

    the total cost

**Hint:** You may want to use the list comprehension function.


Example:

    --Input: 
        N = 4
        K = 2
        C = [1,3,4,5]

    --Output:
        cost = 4 + 5 + 3 * 2 + 1 * 2 = 17

(Why 17 is the minimum cost in this case? You would always like your group to purchase the most expensive candies first. In this way, you will get the multipliers on candy costs for the cheap candies. In this case, two people in your group first purchase candies with costs 4 and 5. And then they will purchase candies with costs 3 and 1. When purchasing candies with costs 3 and 1, both of persons in your group are repeated customers, so they pay 3\*2 and 1\*2.)



Write your code here

<!-- grade_id: cell-562beecb93f03f23, points: 0, grade: false, solution: true -->
```python
def problem2(N, K, C):
    # sort costs decreasing
    C = sorted(C, reverse=True)
    # variable to store the cost
    cost = 0
    # YOUR CODE HERE
    brought_candy = 0
    Candy_num = len(C)
    people_price = []
    people_turn = 0
    ini = 0
    print("total cost = ",end = "")
    while ini < K:
        people_price.append(1)
        ini += 1
    if N < Candy_num:
        for price in C[-N:]:
            if brought_candy == N:
                break
            cost += price
            print(price+ "+")
            brought_candy += 1
    else:
        while brought_candy < N:
            if people_turn == K:
                people_turn = 0
            cost += C[brought_candy] * people_price[people_turn] 
            print(f"{C[brought_candy]} * {people_price[people_turn]} + ",end ="")
            brought_candy += 1
            people_price[people_turn] += 1
            people_turn += 1
    print(f"\n\ntotal cost = {cost}")
    return cost


```

<!-- grade_id: cell-681343a4d7e22bbd, points: 25, grade: true, solution: false -->
```python
# Test cases for Problem Set 4
import random
random.seed(100)
N=random.randint(1,100) 
K=random.randint(1,100)   
C = [random.randint(1,1000) for i in range(N)]
assert problem2(N, K, C) == 8838
```

Output:
```
total cost = 975 * 1 + 820 * 1 + 790 * 1 + 755 * 1 + 750 * 1 + 723 * 1 + 546 * 1 + 520 * 1 + 467 * 1 + 466 * 1 + 444 * 1 + 403 * 1 + 359 * 1 + 270 * 1 + 179 * 1 + 125 * 1 + 113 * 1 + 83 * 1 + 50 * 1 + 

total cost = 8838
```

<!-- grade_id: cell-402bf89f76c4bd42, points: 25, grade: true, solution: false -->
```python
import random
random.seed(10)
N=random.randint(1,100) 
K=random.randint(1,100)   
C = [random.randint(1,1000) for i in range(N)]
assert problem2(N, K, C) == 193878
```

Output:
```
total cost = 993 * 1 + 976 * 1 + 976 * 1 + 964 * 1 + 932 * 1 + 920 * 2 + 881 * 2 + 881 * 2 + 847 * 2 + 844 * 2 + 833 * 3 + 831 * 3 + 791 * 3 + 787 * 3 + 764 * 3 + 703 * 4 + 692 * 4 + 683 * 4 + 678 * 4 + 670 * 4 + 669 * 5 + 629 * 5 + 618 * 5 + 599 * 5 + 597 * 5 + 592 * 6 + 563 * 6 + 550 * 6 + 534 * 6 + 504 * 6 + 502 * 7 + 495 * 7 + 482 * 7 + 474 * 7 + 468 * 7 + 468 * 8 + 462 * 8 + 451 * 8 + 447 * 8 + 440 * 8 + 432 * 9 + 431 * 9 + 391 * 9 + 385 * 9 + 375 * 9 + 372 * 10 + 370 * 10 + 364 * 10 + 336 * 10 + 333 * 10 + 322 * 11 + 311 * 11 + 311 * 11 + 291 * 11 + 285 * 11 + 269 * 12 + 256 * 12 + 246 * 12 + 246 * 12 + 242 * 12 + 212 * 13 + 200 * 13 + 179 * 13 + 165 * 13 + 143 * 13 + 138 * 14 + 137 * 14 + 78 * 14 + 67 * 14 + 46 * 14 + 46 * 15 + 36 * 15 + 16 * 15 + 5 * 15 + 

total cost = 193878
```

```python

```

```python

```

```python

```

Write ChatGPT code here

```python
def problem2(N, K, C):
    # 将糖果价格从高到低排序
    C.sort(reverse=True)
    
    total_cost = 0
    # 每个人买糖果的次数
    purchase_count = [0] * K
    
    for i in range(N):
        # 分配最贵的糖果给当前购买次数最少的人
        person = i % K
        total_cost += (purchase_count[person] + 1) * C[i]
        # 更新该人的购买次数
        purchase_count[person] += 1
        print(purchase_count[person])
    
    return total_cost

# 示例输入
N = 4
K = 2
C = [1, 3, 4, 5]

# 输出总成本
print(min_candy_cost(N, K, C))  # 输出应该为17

```

Output:
```
17
```

```python

```

```python

```

```python

```

Please compare the difference of two codes above. What do you learn after using ChatGPT and how can improve your code? Please write below.

```python
ChatGPT code is more readable, and the i%K and i in range(N) improve the efficiency of the code a lot.
```

```python

```

```python

```